```python
# Initialize Otter
import otter
grader = otter.Notebook("test.ipynb")
```

```python
import pandas as pd
import numpy as np
import matplotlib.pylab as plt
%matplotlib inline
from datascience import *
path="https://raw.githubusercontent.com/fundatosudea/material_20242/develop_otter"
```

<!-- BEGIN QUESTION -->

Ingresa tu nombre  con apellido y número de carnet en las variables dadas más abajo:
Ejemplo:
```python
NombreApellidos="Lisa Simpson"
NumeroCarnet="27182818"
```

```python
NombreApellidos="..."
NumeroCarnet=".."
Email="..."
```

```python
grader.check("qt")
```

Output:
```
qt results:
    qt - 1 result:
        ❌ Test case failed
        Trying:
            assert isinstance(NombreApellidos, str)
        Expecting nothing
        ok
        Trying:
            assert isinstance(NumeroCarnet, str)
        Expecting nothing
        ok
        Trying:
            assert isinstance(Email, str)
        Expecting nothing
        ok
        Trying:
            assert len(NombreApellidos) > 3
        Expecting nothing
        **********************************************************************
        Line 4, in qt 0
        Failed example:
            assert len(NombreApellidos) > 3
        Exception raised:
            Traceback (most recent call last):
              File "/Users/neuman/anaconda3/envs/data8/lib/python3.10/doctest.py", line 1350, in __run
                exec(compile(example.source, filename, "single",
              File "<doctest qt 0[3]>", line 1, in <module>
                assert len(NombreApellidos) > 3
            AssertionError
       
```

<!-- END QUESTION -->

P1. Esta es una pregunta de pruena


```python
characters_q1 = ...
```

```python
grader.check("q1")
```

Output:
```
q1 results:
    q1 - 1 result:
        ❌ Test case failed
        Trying:
            assert characters_q1 == 5
        Expecting nothing
        **********************************************************************
        Line 1, in q1 0
        Failed example:
            assert characters_q1 == 5
        Exception raised:
            Traceback (most recent call last):
              File "/Users/neuman/anaconda3/envs/data8/lib/python3.10/doctest.py", line 1350, in __run
                exec(compile(example.source, filename, "single",
              File "<doctest q1 0[0]>", line 1, in <module>
                assert characters_q1 == 5
            AssertionError
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

These are some submission instructions.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Output:
```
Running your submission against local test cases...


Your submission received the following results when run against available test cases:

    qt results:
        qt - 1 result:
            ❌ Test case failed
            Trying:
                assert isinstance(NombreApellidos, str)
            Expecting nothing
            ok
            Trying:
                assert isinstance(NumeroCarnet, str)
            Expecting nothing
            ok
            Trying:
                assert isinstance(Email, str)
            Expecting nothing
            ok
            Trying:
                assert len(NombreApellidos) > 3
            Expecting nothing
            **********************************************************************
            Line 4, in qt 0
            Failed example:
                assert len(NombreApellidos) > 3
            Exception raised:
                Traceback (most recent call last):
                  File "/Users/neuman/anaconda3/envs/data8/lib/python3.10
```

Tests for q1 (points: 1)
```
>>> assert characters_q1 == 5
```


Tests for qt (points: 0)
```
>>> assert isinstance(NombreApellidos, str)
>>> assert isinstance(NumeroCarnet, str)
>>> assert isinstance(Email, str)
>>> assert len(NombreApellidos) > 3
>>> assert len(NumeroCarnet) > 3
>>> assert len(Email) > 3
```
