```python
# Initialize Otter
import otter
grader = otter.Notebook("test.ipynb")
```

```python
import pandas as pd
import numpy as np
import matplotlib.pylab as plt
%matplotlib inline
from datascience import *
path="https://raw.githubusercontent.com/fundatosudea/material_20242/develop_otter"
```

<!-- BEGIN QUESTION -->

Ingresa tu nombre  con apellido y número de carnet en las variables dadas más abajo:
Ejemplo:
```python
NombreApellidos="Lisa Simpson"
NumeroCarnet="27182818"
```

```python
# BEGIN SOLUTION NO PROMPT
NombreApellidos="Homero Simpson"
NumeroCarnet="27182818"
Email="homero.simpson@udea.edu.co"
# END SOLUTION
""" # BEGIN PROMPT
NombreApellidos="..."
NumeroCarnet=".."
Email="..."
""" # END PROMPT
```

Output:
```
' # BEGIN PROMPT\nNombreApellidos="..."\nNumeroCarnet=".."\nEmail="..."\n'
```

```python
grader.check("qt")
```

<!-- END QUESTION -->

P1. Esta es una pregunta de pruena


```python
# BEGIN SOLUTION NO PROMPT
characters_q1 = 5
# END SOLUTION
""" # BEGIN PROMPT
characters_q1 = ...
""" # END PROMPT
```

Output:
```
' # BEGIN PROMPT\ncharacters_q1 = ...\n'
```

```python
grader.check("q1")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

These are some submission instructions.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1 (points: 1)
```
>>> assert characters_q1 == 5
```


Tests for qt (points: 0)
```
>>> assert isinstance(NombreApellidos, str)
>>> assert isinstance(NumeroCarnet, str)
>>> assert isinstance(Email, str)
>>> assert len(NombreApellidos) > 3
>>> assert len(NumeroCarnet) > 3
>>> assert len(Email) > 3
```
